```python
# Initialize Otter
import otter
grader = otter.Notebook("dsa_summerschool_check.ipynb")
```

## DSA Summer school admittance check

Thanks for your interest in attending DSA Arusha 2022. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Arusha!

```python
# Check the IPython version
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

```python
# Import the good stuff
from nose.tools import assert_equal
import numpy as np
import math
import otter
grader = otter.Notebook()
```

**Question 1:**  Write a function `ypoly` to output the value of `y` given `x` and `p`. Function definition is `y = 4 + 5x^p` Ensure that if the `x` is not positive the value of `y` should be `0`

```python
def ypoly(x, p):
    if x < 0:
        return 0
    y = 4 + 5*x**p
    return y 
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:**  Write a function `perfectly_divisible` to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $q$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

```python
def perfectly_divisible(N, q):
    if N < 0:
        N = 20
    i=0
    liste = []
    while (q*i<=N):
        liste.append(q*i)
        i+=1
    
    return liste
        
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:**  Write a function `flatten_lists` that takes in a list of lists and outputs a sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_lists(superlist):
    liste =[]
    for i in superlist:
        liste =liste + i
    liste.sort()
    return liste
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:**  Write a function `longest_word` to return the longest string in a line of text <br>
Example: for text `Hallo people of the world!` should return `people` <br>
If there is a tie, then it should return the `1st` "longest" word encountered

```python
def longest_word(sentence):
    liste = sentence.split(" ")
    l = sorted(liste, key = len, reverse = True)  
    return l[0]
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5:**  Write a function `dictionarize` that takes a list as input and returns a dictionary with the elements of the list as the keys and the frequency of the element as the values.

```python
def dictionarize(alist):
    d = {}
    for i in alist:
        value =d.get(i,0)
        d[i]=value+1
        
    return d
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a,b):
    if a==0:
        return b
    
    return greatest_common_divisor(b%a,a)
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:**  Write a function `get_nearest_farthest` that takes in a point of interest (pt) and a list of points and returns the nearest point and the farthest point from the point of intest.

```python
def get_nearest_farthest(pt, ptlist):
    p=q=ptlist[0]
    for point in ptlist:
        if math.dist(pt,point) < math.dist(p,pt):
            p=point
        if math.dist(pt,q) < math.dist(point,pt):
            q=point
    return p,q
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload the generated zip file to the Google sheet please. Uploading work which is not yours will lead to non admittance.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Output:
```
There was an error generating your LaTeX; showing full error message:
================================================================================
This is BibTeX, Version 0.99d (TeX Live 2020/Debian)
The top-level auxiliary file: notebook.aux
I found no \citation commands---while reading file notebook.aux
I found no \bibdata command---while reading file notebook.aux
I found no \bibstyle command---while reading file notebook.aux
(There were 3 error messages)

================================================================================
If the error above is related to xeCJK or fandol in LaTeX and you don't require this functionality, try running again with no_xecjk set to True or the --no-xecjk flag.
/home/hagbe/.local/lib/python3.9/site-packages/otter/check/notebook.py:330: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
<IPython.core.display.HTML object>
```

Tests for q1 (points: 6)
```
>>> ypoly(3,3) == 139
True
```
```
>>> ypoly(-2,3) == 0
True
```


Tests for q2 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q3 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```


Tests for q4 (points: 5)
```
>>> assert_equal(longest_word("Hello there people"), 'people')
```


Tests for q5 (points: 5)
```
>>> assert_equal(dictionarize([1,2,2,3,4,5]), {1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
```


Tests for q6 (points: 5)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```
